```python
# Initialize Otter
import otter
grader = otter.Notebook("hw02.ipynb")
```

<div class="alert alert-success" markdown="1">

#### Homework 2

# Arrays and DataFrames

### EECS 398-003: Practical Data Science, Fall 2024

#### Due Thursday, September 12th at 11:59PM
    
</div>

## Instructions

Welcome to Homework 2! In this homework, you will practice with `numpy` arrays and `pandas` DataFrames, as we saw in Lectures 3 and 4. See the [Readings section of the Resources tab on the course website](https://practicaldsc.org/resources/#readings) for supplemental resources.

You are given six slip days throughout the semester to extend deadlines. See the [Syllabus](https://practicaldsc.org/syllabus) for more details. With the exception of using slip days, late work will not be accepted unless you have made special arrangements with your instructor.

To access this notebook, you'll need to clone our [public GitHub repository](https://github.com/practicaldsc/fa24/). The [⚙️ Environment Setup](https://practicaldsc.org/env-setup) page on the course website walks you through the necessary steps. Once you're done, you'll submit your completed notebook to Gradescope.

Please start early and submit often. You can submit as many times as you'd like to Gradescope, and we'll grade your **most recent** submission. Remember that the public `grader.check` tests in your notebook are not comprehensive, and that your work will also be graded on hidden test cases on Gradescope after the submission deadline.

This homework is worth a total of **51 points**, 49 of which come from the autograder, **and 2 of which are manually graded by us** (Question 3). The number of points each question is worth is listed at the start of each question. **The three parts of the assignment are independent, so feel free to move around if you get stuck**. Tip: if you're using Jupyter Lab, you can see a Table of Contents for the notebook by going to View > Table of Contents.

To get started, run the **two** import cells below, plus the cell at the top of the notebook that imports and initializes `otter`. The first cell below installs a new package that wasn't included in the `pds` conda environment that we'll need for Part 1.

```python
!pip install networkx
```

```python
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
import seaborn as sns
from matplotlib_inline.backend_inline import set_matplotlib_formats

import plotly
import plotly.graph_objects as go
import plotly.express as px
import plotly.io as pio

# Preferred styles
pio.templates["pds"] = go.layout.Template(
    layout=dict(
        margin=dict(l=30, r=30, t=30, b=30),
        autosize=True,
        width=600,
        height=400,
        xaxis=dict(showgrid=True),
        yaxis=dict(showgrid=True),
        title=dict(x=0.5, xanchor="center"),
    )
)
pio.templates.default = "simple_white+pds"

set_matplotlib_formats("svg")
sns.set_context("poster")
sns.set_style("whitegrid")
plt.rcParams["figure.figsize"] = (10, 5)

# Use plotly as default plotting engine
pd.options.plotting.backend = "plotly"

import networkx as nx
```

## Part 1: PageRank 🔗

---

In this part of the homework, you'll replicate the PageRank algorithm, the algorithm that Google uses to decide how to rank search results. Along the way, you'll develop proficiency with using arrays in the context of linear algebra.

### Overview

Larry Page and Sergey Brin developed the Google search engine (originally known as "BackRub") while PhD students at Stanford University. Their key innovation was the **PageRank** algorithm. PageRank works by assigning each page on the internet a "score" based on its relative importance. Today, PageRank is one of the many algorithms that Google uses in determining how to rank search results.

We'll give an explanation of the necessary components of PageRank here, but you may also want to refer to these slides, which inspired our exposition: [**The Billion Dollar Eigenvector**](https://jdc.math.uwo.ca/M1600b-2014/l/pagerank-1600.pdf).

**Key idea: The more incoming links a page has, the more important it is!**

<center>

<img src="imgs/network.png" width=300><br><small>An example network of pages.</small>
    
</center>

There are a few naïve approaches to scoring pages given that key idea, but they have their flaws:
- One idea: the score of a page is equal to the number of pages that link to it. This doesn't work because it doesn't account for the importance of the incoming links. For example, if my personal website links to [apple.com](https://apple.com), that doesn't make [apple.com](https://apple.com) any more important than it already was. But if [apple.com](https://apple.com) links to my personal website, that must mean my personal website is important!
- Another idea: The score of a page is equal to the sum of the scores of the pages that link to it. This also doesn't work, because there is no non-zero solution!

The actual solution assigns scores to pages such that **when a page links to several other pages, the score it gives to them is shared**.

In the example graph above, we see that:
- Page 1 links to Page 2.
- Page 2 links to Pages 1 and 4.
- Page 3 links to Pages 1 and 4.
- Page 4 links to Pages 1, 2, and 3.

If we let $x_i$ be the score assigned to Page $i$, we're looking for the solution to the equations:

$$\begin{align*}
x_1 &= &\frac{1}{2}x_2 &+ \frac{1}{2}x_3 &+ \frac{1}{3}x_4 \\
x_2 &= x_1 &&           &+ \frac{1}{3}x_4 \\
x_3 &= &&               & \frac{1}{3}x_4 \\
x_4 &= &\frac{1}{2}x_2 &+ \frac{1}{2}x_3
\end{align*}$$

A more concise way to express the system above is in matrix form, which we can do like so:

$$\vec{x} = A \vec{x}, \qquad A = \begin{bmatrix}
0 & \frac{1}{2} & \frac{1}{2} & \frac{1}{3} \\
1 & 0 & 0 & \frac{1}{3} \\
0 & 0 & 0 & \frac{1}{3} \\
0 & \frac{1}{2} & \frac{1}{2} & 0
\end{bmatrix} \qquad \vec x = \begin{bmatrix} x_1 \\ x_2 \\ x_3 \\ x_4 \end{bmatrix}$$

**Our goal**: Given an adjacency matrix $A$, find the solution to the equation $\vec{x} = A \vec{x}$. We will return to solving this goal in Question 2, but keep it at the back of your mind.

For now, we'll aim to understand the structure of adjacency matrices.
- Each **column** of an adjacency matrix describes the movement **from** a given page; the sums of the columns in an adjacency matrix are equal to 1.
- Each **row** of an adjacency matrix describes the movement **into** a given page; the sums of the rows in an adjacency matrix don't necessarily add to 1.

Below, we've defined a 2D array, `A`, representing $A$ above:

```python
A = np.array([[0, 1 / 2, 1 / 2, 1 / 3],
              [1, 0, 0, 1 / 3],
              [0, 0, 0, 1 / 3],
              [0, 1 / 2, 1 / 2, 0]])
A
```

The function below draws a network given an adjacency matrix. Run the next two cells.

```python
def plot_from_adjacency(adjacency_matrix, node_sizes=0.25):
    np.random.seed(25)
    plt.figure(figsize=(8, 5))
    G = nx.from_numpy_array(adjacency_matrix.T, create_using=nx.DiGraph)
    layout = nx.spring_layout(G)
    labels_dict={i: i+1 for i in range(adjacency_matrix.shape[0])}
    nx.draw(G, layout, 
            node_size=15000 * node_sizes, labels=labels_dict, with_labels=True, font_color='white', font_weight='bold', font_size=15, 
            connectionstyle='arc3, rad = 0.1')
    plt.show()

plot_from_adjacency(A)
```

Cool!

### Question 1: Creating Adjacency Matrices  <div style="display:inline-block; vertical-align: middle; padding:7px 7px; font-size:10px; font-weight:light; color:white; background-color:#e84c4a; border-radius:7px; text-align:left;">4 Points</div>

Having to specify an adjacency matrix manually is slightly cumbersome. It is more natural and convenient for us to describe the links between webpages using a dictionary. Once such example is as follows:

```python
example_net = {
    1: [2],
    2: [1, 4],
    3: [1, 4],
    4: [1, 2, 3]
}
```

In the above "network dictionary", we are told that:
- Page 1 links to Page 2.
- Page 2 links to Pages 1 and 4.
- Page 3 links to Pages 1 and 4.
- Page 4 links to Pages 1, 2, and 3.

**Note that this dictionary describes the same network that the adjacency matrix `A` does.**

Below, complete the implementation of the function `create_adjacency`, which takes in a network dictionary, `network` (formatted similarly to `example_net`) and returns the adjacency matrix that corresponds to `network`. Example behavior is given below.

```python
>>> create_adjacency(example_net)
array([[0.        , 0.5       , 0.5       , 0.33333333],
       [1.        , 0.        , 0.        , 0.33333333],
       [0.        , 0.        , 0.        , 0.33333333],
       [0.        , 0.5       , 0.5       , 0.        ]])
```

A few notes:
- It is **not** guaranteed that there are 4 pages in the network. 
- It **is** guaranteed that all pages link to at least one page – potentially itself – and that there is at least one page.
- Remember that adjacency matrices are 1-indexed, like in math, but Python is 0-indexed.

***Hint***: Look into `np.zeros`. You _can_ use a (nested) `for`-loop.

```python
def create_adjacency(network):
    ...

# Feel free to change this input to make sure your function works correctly.
create_adjacency(example_net)
```

```python
grader.check("q01")
```

### Question 2: Ranking Pages

Now, back to the task at hand.

**Our goal**: Given an adjacency matrix $A$, find the solution to the equation $\vec{x} = A \vec{x}$. In our example network $A$, the setup looks like:

$$A = \begin{bmatrix}
0 & \frac{1}{2} & \frac{1}{2} & \frac{1}{3} \\
1 & 0 & 0 & \frac{1}{3} \\
0 & 0 & 0 & \frac{1}{3} \\
0 & \frac{1}{2} & \frac{1}{2} & 0
\end{bmatrix} \qquad \vec x = \begin{bmatrix} x_1 \\ x_2 \\ x_3 \\ x_4 \end{bmatrix}$$

How do we find $\vec x$? It turns out that for all valid adjacency matrices, $\vec x$ is an **eigenvector** that corresponds to the **eigenvalue of 1**. Not sure what that means? That's okay – you can learn more about eigen-everything [here](https://youtu.be/PhfbEr2btGQ?si=1diIM4b31qSYrDOZ).

But for our purposes, you don't actually _need_ to know about eigenvalues and eigenvectors. That's because, due to some mathematical theory that's outside of the scope of the course, we can estimate $\vec x$ with great accuracy if we follow the following process:

1. Start by initializing $\vec x^{(0)}$ to a uniform array that sums to 1. This is our initial, or 0th, guess of the true value of $\vec x$. In our example $A$, since $\vec x \in \mathbb{R}^4$, we'd initialize $\vec x^{(0)} = \begin{bmatrix} \frac{1}{4} \\ \frac{1}{4} \\ \frac{1}{4} \\ \frac{1}{4} \end{bmatrix}$. You can imagine this means that a user is equally likely to be on any of the four pages before clicking anything.
1. Simulate the process of the user clicking one link by multiplying $A$ by $\vec x^{(0)}$, and call this $\vec x^{(1)}$, our next guess of the true $\vec x$. In other words, $\vec x^{(1)} = A \vec x^{(0)}$.
2. Repeat this again, so $\vec x^{(2)} = A \vec x^{(1)} = A^2 \vec x^{(0)}$.
3. Repeat many, many times, until the difference between $\vec x^{(i)}$ and $\vec x^{(i-1)}$ is minimal.

Eventually, our guesses $\vec x^{(i)}$ will converge on the true value of $\vec x$, which we can interpret as containing the probability that a user is on any particular page in the long run. **Larger probabilities mean more important pages!**

In short, to estimate $\vec x$, we evaluate $A^{m} \vec x^{(0)}$, where $m$ is some large integer and $x^{(0)}$ is a uniform vector. $m = 100$ is usually sufficient.

Let's work this out concretely with our example adjacency matrix, $A$. Let's initialize `x` to an array of length 4 in which each value is `1 / 4`.

```python
x = np.ones(4) / 4
x
```

A better solution than writing:

```python
A @ A @ A @ A @ A @ A @ ... @ x
```

or even:

```python
for i in range(100):
    x = A @ x
```

is using `np.linalg.matrix_power`:

```python
scores = np.linalg.matrix_power(A, 100) @ x
scores
```

Above are what we call the PageRank scores of each page. Remember, larger scores are more important! This is telling us that:
- Page 2 is the most important,
- then Page 1,
- then Page 4,
- then Page 3.

So, if we were Google, we'd show Page 2 first, then Page 1, then Page 4, then Page 3.

Now it's your turn! We've broken this question up into two smaller parts.

#### Question 2.1: Computing Scores  <div style="display:inline-block; vertical-align: middle; padding:7px 7px; font-size:10px; font-weight:light; color:white; background-color:#e84c4a; border-radius:7px; text-align:left;">2 Points</div>

Complete the implementation of the function `compute_scores`, which takes in an adjacency matrix `matrix` of shape `(n, n)`, where `n` is a positive integer, and returns an array of length `n` containing the PageRank scores of all pages. Compute the PageRank scores using a matrix power of 100, as we did above. Example behavior is given below.

```python
# Remember, compute_scores should work for any adjacency matrix, not just A!
>>> compute_scores(A)
array([0.30769231, 0.38461538, 0.07692308, 0.23076923])
```

```python
def compute_scores(matrix):
    ...

# Feel free to change this input to make sure your function works correctly.
compute_scores(A)
```

```python
grader.check("q02_1")
```

We can change the sizes of the pages in our network to be proportional to their PageRank scores. To do this, use the `node_sizes` argument in `plot_from_adjacency`.

```python
plot_from_adjacency(A, node_sizes=compute_scores(A))
```

#### Question 2.2: Ranking Pages  <div style="display:inline-block; vertical-align: middle; padding:7px 7px; font-size:10px; font-weight:light; color:white; background-color:#e84c4a; border-radius:7px; text-align:left;">2 Points</div>

Complete the implementation of the function `pagerank`, which takes in an adjacency matrix `matrix` of shape `(n, n)`, where `n` is a positive integer, and returns a list containing the **numbers** of the pages in the matrix, in **decreasing** order of PageRank score as computed by `compute_scores`. If there are ties in scores, return the pages in any order. Example behavior is given below.

```python
>>> pagerank(A)
[2, 1, 4, 3]
```

***Hint:*** Look into `np.argsort`. This should only take 2-3 lines; do not write a `for`-loop.

```python
def pagerank(matrix):
    ...

# Feel free to change this input to make sure your function works correctly.
pagerank(A)
```

```python
grader.check("q02_2")
```

Once you've completed `pagerank`, run the following cell to try your work .

```python
# Here's another example network, which you can use to test your code.
test_net = {1: [6], 2: [1, 3, 4, 6], 3: [2, 5, 6], 4: [1], 5: [1, 2, 6], 6: [2, 4]}
print(test_net)
test_net_adjacency = create_adjacency(test_net)
test_net_scores = compute_scores(test_net_adjacency)
plot_from_adjacency(test_net_adjacency, node_sizes=test_net_scores)
```

Great work! You've recreated the magic behind Google. (Does this mean you automatically get a job at Google now? Unfortunately, no...)

### Question 3: Sinkholes  <div style="display:inline-block; vertical-align: middle; padding:7px 7px; font-size:10px; font-weight:light; color:white; background-color:#e84c4a; border-radius:7px; text-align:left;">2 Points</div>

Let's wrap up this part of the homework with a puzzle. Consider the following network:

```python
weird_net = {
    1: [1],
    2: [1, 3, 5],
    3: [2, 4],
    4: [1, 2, 3],
    5: [3]
}
```

Note that Page 1 links to itself and to no other pages. Practically speaking, we can interpret Page 1 as being a "dead end" or "sinkhole", with no outgoing links.

Run the cells below to compute the adjacency matrix and scores for the above network and to visualize it.

```python
weird_matrix = create_adjacency(weird_net)
weird_scores = compute_scores(weird_matrix)
print(weird_scores)
plot_from_adjacency(weird_matrix, node_sizes=weird_scores)
```

<!-- BEGIN QUESTION -->

It appears that Page 1's score is 1, and all of the other pages' scores are 0! (`9.6409e-10` means $9.64 \cdot 10^{-10}$, which is a really, really small number.)

<span style="color:red"><b>In the cell below, write your answers to the following two _manually-graded_ prompts.</b></span>

- Why do you think Page 1's score is so high, and the other pages' scores are so low? (***Hint:*** Think about how we interpret the score of a page.)
- Read the [Damping factor](https://en.wikipedia.org/wiki/PageRank#Damping_factor) section of the Wikipedia article on PageRank. In two sentences, describe (to the best of your ability) how using damping would prevent the score of Page 1 from becoming 1. (If you read the article closely, the answer is there – describe it in your own words.)

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Part 2: Euchre Returns! 🃏

---

In this part of the homework, we'll practice simulating various probabilities in the context of the popular card game, [Euchre](https://www.wikihow.com/Play-Euchre)!

Euchre, as you may have seen in EECS 280, is a 4-player card game, in which the players are named Player 0, Player 1, Player 2, and Player 3. Euchre is played with a deck of **24 cards**. Each card has two attributes:
- a **suit**, which is either hearts (<span style="color:red">❤️</span>), diamonds (<span style="color:red">♦️</span>), clubs (♣), or spades (♠).
- a **face value**, which is either 9, 10, Jack (J), Queen (Q), King (K), or Ace (A).

The full deck of 24 cards is shown below.

<div align="center">

<span style="color:red">❤️: 9, 10, J, Q, K, A</span><br>
<span style="color:red">♦️: 9, 10, J, Q, K, A</span><br>
<span style="color:black">♣: 9, 10, J, Q, K, A</span><br>
<span style="color:black">♠: 9, 10, J, Q, K, A</span><br>

</center>

The 24 cards are all different. The list `deck`, defined below, represents our deck of 24 cards. Each element in deck is a string representing a card, in the format `'{face} {suit}'`.

Note that you don't need to know how Euchre is played, beyond what is explained above, to work on the following question. (Suraj has never played Euchre before – he's never taken EECS 280 either!)

```python
# Do NOT edit this cell! 
# In our hidden tests, we will assume deck is defined exactly as below.
suits = ['Hearts', 'Diamonds', 'Clubs', 'Spades']
faces = ['9', '10', 'J', 'Q', 'K', 'A']
deck = [f'{face} {suit}' for suit in suits for face in faces]
deck
```

At the start of a round of Euchre, the 24 cards are randomly distributed to the 4 players such that each player gets 5 cards, and the remaining 4 cards ($24 - 5 \cdot 4 = 4$) are put to the side. 

You should use functions built into `np.random` to simulate the act of shuffling and distributing cards. You may find `np.random.choice` and `np.random.permutation` to be useful, but be careful about the arguments you provide and the assumptions they make. You're free to define helper functions to use throughout your work, too – **we did this ourselves** – and you'll need to research various string, list, array, and `np.random` functions and methods to get this all to work.

Your answers will be (very) slightly different than those of your peers, and slightly different each time you run your notebook, due to randomness. This is expected, and our autograder tests will account for this. **Don't** work out the math by hand – that isn't the point of these questions, rather the point is to build your fluency in simulating.

### Question 4: At Least $k$ Same Suit  <div style="display:inline-block; vertical-align: middle; padding:7px 7px; font-size:10px; font-weight:light; color:white; background-color:#e84c4a; border-radius:7px; text-align:left;">5 Points</div>

Complete the implementation of the function `prob_at_least_k_same_suit`, which takes in an integer `k` and returns a simulated estimate of the **probability that Player 0 ends up with at least $k$ cards of the same suit**. Remember, the 24 cards in the Euchre deck (`deck`) are distributed such that Player 0, Player 1, Player 2, and Player 3 each receive 5 cards, and 4 cards are put to the side.

Example behavior is given below.

```python
>>> prob_at_least_k_same_suit(4)
0.02693

# Each time we call prob_at_least_k_same_suit with the same input,
# we should see a slightly different result.
>>> prob_at_least_k_same_suit(4)
0.02591
```

Some extra guidance:
- Use 100,000 repetitions in your simulation.
- If `k` is negative or `0`, return `0`. If `k` is greater than the number of cards given to Player 0, return `0`. (Note that there are some values of `k` for which you'll always end up returning `1` – it's useful to think about which values of `k` these are. Think back to the pigeonhole principle from EECS 203!)
- Look at the default `replace` argument in `np.random.choice` – you may need to set this to something else.

```python
def prob_at_least_k_same_suit(k):
    ...

# Feel free to change this input to make sure your function works correctly.
# Each time you run this cell, you should see a slightly different result!
prob_at_least_k_same_suit(4)
```

```python
grader.check("q04")
```

### Question 5: Evenly Distributed Jacks  <div style="display:inline-block; vertical-align: middle; padding:7px 7px; font-size:10px; font-weight:light; color:white; background-color:#e84c4a; border-radius:7px; text-align:left;">5 Points</div>

Complete the implementation of the function `prob_all_one_jack`, which takes no arguments and returns a simulated estimate of the **probability that each player ends up with exactly one Jack (J)**.

Again, use 100,000 repetitions in your simulation.

```python
def prob_all_one_jack():
    ...

# Each time you run this cell, you should see a slightly different result!
prob_all_one_jack()
```

```python
grader.check("q05")
```

## Part 3: Movies 🎥

---

In this final part of the homework, we'll build our skills in working with DataFrames in `pandas`.


Run the cell below to load in a dataset with information about various movies from [IMDb](https://www.imdb.com/), the Internet Movie Database. (The dataset comes from [Kaggle](https://www.kaggle.com/datasets/parthdande/imdb-dataset-2024-updated)).

```python
imdb = pd.read_csv('data/imdb-2024-cleaned.csv')
imdb
```

As the preview above shows us, `imdb` has 255 rows, and has some brand new movies, like `'A Quiet Place: Day One'`, released earlier this year.

Not sure what one of the columns means? Google it, as data scientists do – you'll find helpful information directly on IMDb's website.

In lecture, we learned that it's good practice to set the index of a DataFrame to a unique identifier for each row, if one exists. At first glance, it may seem like the values in the `'Title'` column are unique, but upon further investigation we see that there are actually duplicate `'Title'`s:

```python
imdb['Title'].value_counts()
```

There are 2 rows for `'Planet of the Apes'`, for example. We can query to see just those rows:

```python
imdb[imdb['Title'] == 'Planet of the Apes']
```

It seems like there are indeed multiple different `'Planet of the Apes'` movies, released in different years. So, for now, we'll leave the index of `imdb` as-is.

Time to get to work!  Note that all of the following questions are answerable using concepts from [Lecture 4](https://practicaldsc.org/resources/lectures/lec04/lec04-filled.html), though you may find them to be slightly easier after you've gained more practice in Lecture 5.

Important: **DO NOT** modify the `imdb` DataFrame at any point in this notebook! If you do, your code won't be graded correctly when you submit it to Gradescope.

### Question 6: Exploring the Data

In this question, which is made up of 7 smaller subparts, we'll ask you to answer various questions involving the `imdb` DataFrame. **Don't** hard-code your answers; use `pandas` code to find them programatically.

#### Question 6.1: Highest Rating Ever  <div style="display:inline-block; vertical-align: middle; padding:7px 7px; font-size:10px; font-weight:light; color:white; background-color:#e84c4a; border-radius:7px; text-align:left;">1 Point</div>

Assign `highest_rating_ever` to the highest `'IMDb Rating'` of any movie in `imdb`. Your answer should be a float.

```python
highest_rating_ever = ...
highest_rating_ever
```

```python
grader.check("q06_1")
```

#### Question 6.2: Movie with Highest Rating Ever  <div style="display:inline-block; vertical-align: middle; padding:7px 7px; font-size:10px; font-weight:light; color:white; background-color:#e84c4a; border-radius:7px; text-align:left;">2 Points</div>

Assign `movie_with_highest_rating_ever` to the `'Title'` of the movie with the highest `'IMDb Rating'` of any movie in `imdb`. Assume there are no ties. Your answer should be a string.

```python
movie_with_highest_rating_ever = ...
movie_with_highest_rating_ever
```

```python
grader.check("q06_2")
```

#### Question 6.3: Comedy Movies  <div style="display:inline-block; vertical-align: middle; padding:7px 7px; font-size:10px; font-weight:light; color:white; background-color:#e84c4a; border-radius:7px; text-align:left;">1 Point</div>

Assign `num_comedy` to the number of movies in `imdb` with a `'Genre'` of `'Comedy'`. Your answer should be an int.

```python
num_comedy = ...
num_comedy
```

```python
grader.check("q06_3")
```

#### Question 6.4: Most Common Genre in 2024  <div style="display:inline-block; vertical-align: middle; padding:7px 7px; font-size:10px; font-weight:light; color:white; background-color:#e84c4a; border-radius:7px; text-align:left;">2 Points</div>

Assign `most_common_genre_2024` to the name of the most common `'Genre'` among all movies in `imdb` released in 2024. Assume there are no ties. Your answer should be a string.

```python
most_common_genre_2024 = ...
most_common_genre_2024
```

```python
grader.check("q06_4")
```

#### Question 6.5: Big Directors <div style="display:inline-block; vertical-align: middle; padding:7px 7px; font-size:10px; font-weight:light; color:white; background-color:#e84c4a; border-radius:7px; text-align:left;">3 Points</div>

Assign `prop_big_directors` to the **proportion** of movies in `imdb` directed by the top 4 `'Director'`s, when `'Director'`s are ranked by the number of movies they've directed in descending order. Assume that the top 4 `'Directors'` are unambiguous; i.e. that there isn't a 5th director that is tied with one of the top 4. Your answer should be a float between 0 and 1.

```python
prop_big_directors = ...
prop_big_directors
```

```python
grader.check("q06_5")
```

#### Question 6.6: Harry Potter  <div style="display:inline-block; vertical-align: middle; padding:7px 7px; font-size:10px; font-weight:light; color:white; background-color:#e84c4a; border-radius:7px; text-align:left;">2 Points</div>

Assign `average_harry_potter` to the mean `'IMDb Rating'` of all movies with the string `'Harry Potter'` in the `'Title'`. Your answer should be a float.

```python
average_harry_potter = ...
average_harry_potter
```

```python
grader.check("q07_6")
```

#### Question 6.7: Family Friendly  <div style="display:inline-block; vertical-align: middle; padding:7px 7px; font-size:10px; font-weight:light; color:white; background-color:#e84c4a; border-radius:7px; text-align:left;">4 Points</div>

Assign `num_hours_family_friendly` to the number of **hours** it would take you to watch the 10 highest-rated `'PG-13'` movies whose `'Genre'` is either `'Action'` or `'Adventure'`. By highest-rated, we're referring to movies with the highest `'IMDb Rating'`s. Assume that there are no ties in `'IMDb Rating'`s. Your answer should be a float.

```python
num_hours_family_friendly =  ...
num_hours_family_friendly
```

```python
grader.check("q06_7")
```

### Question 7: Star Struck

Take a look at the `'Star Cast'` column of `imdb`, which we didn't use in Question 7:

```python
imdb['Star Cast']
```

Right now, actors' names aren't separated by spaces. With `'Star Cast'` in this form, we can still perform _some_ queries, like looking at all of the movies Margot Robbie was in:

```python
imdb[imdb['Star Cast'].str.contains('Margot Robbie')]
```

But we **can't** answer questions like, which actor was in the most movies? Let's change that.

#### Question 7.1: Data Cleaning  <div style="display:inline-block; vertical-align: middle; padding:7px 7px; font-size:10px; font-weight:light; color:white; background-color:#e84c4a; border-radius:7px; text-align:left;">5 Points</div>

Complete the implementation of the function `extract_names`, which takes in `s`, a single string formatted like those in the `'Star Cast'` column of `imdb`, and returns a **list** of strings with the names of all of the actors in `s`. Example behavior is given below.

```python
>>> extract_names('Makoto ShinkaiClark Cheng')
['Makoto Shinkai', 'Clark Cheng']

>>> extract_names('Anya Taylor-JoyChris HemsworthTom Burke')
['Anya Taylor-Joy', 'Chris Hemsworth', 'Tom Burke']

>>> extract_names('Santa Ono')
['Santa Ono']
```

For simplicity, you **must** assume that:
- Each actor only has a single first name and a single last name.
- First names and last names are both at least two letters long, contain no spaces, and contain exactly one uppercase letter
- If there is a dash (`'-'`) in an actor's first name or last name, it is not the first or last letter. Dashes are the only non-alphabetic characters that can appear in a first name or last name.

Note that the above assumptions mean that `extract_names` _will_ work incorrectly for cases like the one below; we're adding this simplifying assumption to make your life easier.

```python
>>> extract_names('Leonardo DiCaprioJonah HillMargot Robbie')
['Leonardo Di', 'Caprio', 'Jonah Hill', 'Margot Robbie']
```

Note that the implementation of `extract_names` doesn't involve any `pandas` methods; it's a pure Python problem.

```python
def extract_names(s):
    ...

# Feel free to change this input to make sure your function works correctly.
extract_names('Makoto ShinkaiClark Cheng')
```

```python
grader.check("q07_1")
```

Now, using syntax we'll see in Week 3, we've defined for you a Series, named `star_names`, containing the names of all individual actors in `'Star Cast'`. This involves using the Series `apply` method to call your `extract_names` function on every value in `imdb['Star Cast']`, and then combining the resulting lists into one massive list, and finally converting that to Series.

```python
star_names = pd.Series(imdb['Star Cast'].apply(extract_names).sum())
star_names
```

Note that the values in `star_names` are not unique!

#### Question 7.2: Modes  <div style="display:inline-block; vertical-align: middle; padding:7px 7px; font-size:10px; font-weight:light; color:white; background-color:#e84c4a; border-radius:7px; text-align:left;">2 Points</div>

Complete the implementation of the function `most_common`, which takes in a Series, `ser`, and returns a **list** containing the mode(s) of `ser`. The order of the elements in the returned list does not matter. Assume `ser` contains at least one element. Example behavior is given below.

```python
>>> most_common(pd.Series([1, 2, 2]))
[2]

>>> most_common(pd.Series([1, 2]))
[1, 2]

# This works strangely because of our assumptions in 7.1.
# We're intentionally not showing you the entire output.
>>> star_names_output = most_common(star_names)
>>> 'Caprio' in star_names_output and 'Leonardo Di' in star_names_output
True
```

```python
def most_common(ser):
    ...

# Feel free to change this input to make sure your function works correctly.
most_common(pd.Series([1, 2, 2]))
```

```python
grader.check("q07_2")
```

After you've implemented `most_common`, run the cell below.

```python
most_common(star_names)
```

You may recognize some of these names, but not others. You should explore! For instance, query `imdb` for all the movies that `'Pete Docter'` starred in. What do you notice?

```python
# You're NOT required to do this, but you should – it's fun!
# BEGIN SOLUTION
imdb[imdb['Star Cast'].str.contains('Pete Docter')]
# END SOLUTION
```

### Question 8: Diving Deeper

To wrap up, let's revisit the larger DataFrame to answer a few more involved questions. Make sure your functions don't make in-place modifications to the input DataFrames!

#### Question 8.1: Longest Title  <div style="display:inline-block; vertical-align: middle; padding:7px 7px; font-size:10px; font-weight:light; color:white; background-color:#e84c4a; border-radius:7px; text-align:left;">3 Points</div>

Complete the implementation of the function `longest_title`, which takes in a DataFrame `df` and returns the longest movie `'Title'`. Assume that `df` has the same 9 column titles as `imdb`, with the same data types, but potentially a different number of rows in a different order, with a potentially different index. Assume that `df` has at least one row, and assume there are no ties. Example behavior is given below.

```python
# Remember, imdb.head(5) is a DataFrame with just the first 5 rows of imdb.
>>> longest_title(imdb.head(5))
'10 Things I Hate About You'
```

```python
def longest_title(df):
    ...

# Feel free to change this input to make sure your function works correctly.
longest_title(imdb.head(5))
```

```python
grader.check("q08_1")
```

Once you've implemented `longest_title`, run the cell below to see the longest movie title in `imdb`.

```python
longest_title(imdb)
```

#### Question 8.2: Outliers  <div style="display:inline-block; vertical-align: middle; padding:7px 7px; font-size:10px; font-weight:light; color:white; background-color:#e84c4a; border-radius:7px; text-align:left;">3 Points</div>

Typically, it seems that `'IMDb Rating'`, which are determined by user ratings, and `'MetaScore'` ratings, which are weighted averages of scores given by trusted movie critics, are correlated. That is, movies with high `'IMDb Rating'`s tend to have high `'MetaScore'` ratings, and movies with low `'IMDb Rating'`s tend to have low `'MetaScore'` ratings. We can see this trend in the following scatter plot. Hover over points to see the names of the movies.

```python
imdb.plot(kind='scatter', x='IMDb Rating', y='MetaScore', hover_name='Title', width=800, height=600)
```

As you can see, there are some outliers, i.e. movies that had relatively high `'MetaScore'`s but relatively low `'IMDb Rating'`s, or vice versa.

We define the `'MetaScore'` to `'IMDb Rating'` ratio of a movie as follows:

$$\frac{\text{Movie's MetaScore}}{\text{Movie's IMDb Rating}}$$

Complete the implementation of the function `metascore_to_rating_outlier`, which takes in a DataFrame `df` and returns the movie with the **largest** `'MetaScore'` to `'IMDb Rating'`. (Note that this is only one of the two types of outliers mentioned in the paragraph at the top of this cell, but use this definition in your solution.)

Assume that `df` has the same 9 column titles as `imdb`, with the same data types, but potentially a different number of rows in a different order, with a potentially different index. Assume that `df` has at least one row, and assume there are no ties. Example behavior is given below.


```python
>>> metascore_to_rating_outlier(imdb.head(5))
'A Quiet Place'
```

```python
def metascore_to_rating_outlier(df):
    ...

# Feel free to change this input to make sure your function works correctly.
metascore_to_rating_outlier(imdb.head(5))
```

```python
grader.check("q08_2")
```

#### Question 8.3: Recommended For Me <div style="display:inline-block; vertical-align: middle; padding:7px 7px; font-size:10px; font-weight:light; color:white; background-color:#e84c4a; border-radius:7px; text-align:left;">3 Points</div>

And finally, now that you're almost done Homework 2, it's time to find new movies in your favorite genre that the data thinks you might like.

Complete the implementation of the function `genre_specific`, which takes in a DataFrame `df` and a string, `genre`, corresponding to a `'Genre'` in `imdb`. `genre_specific(df, genre)` should return a DataFrame that:
- Contains all of the movies in `df` of `'Genre'` `genre` **released in 2024**. Assume `genre` is a valid value in `df['Genre']`.
- Is sorted in descending order of `'IMDb Rating'`, with ties broken by `'MetaScore'` in descending order. Assume no two movies have both the same `'IMDb Rating'` and the same `'MetaScore'`. (There are a few cases where this happens, but you won't be tested on them in the hidden test cases.)
- Is indexed by `'Title'`, and only has two columns: `'IMDb Rating'` and `'MetaScore'`.

Example behavior is given below.

```python
>>> genre_specific(imdb, 'Horror')
```

<div style="text-align: left;">
  <table border="1" class="dataframe">
    <thead>
      <tr>
        <th style="text-align: left;"></th>
        <th style="text-align: left;">IMDb Rating</th>
        <th style="text-align: left;">MetaScore</th>
      </tr>
      <tr>
        <th>Title</th>
        <th></th>
        <th></th>
      </tr>
    </thead>
    <tbody>
      <tr>
        <th>Alien: Romulus</th>
        <td>7.1</td>
        <td>66.9</td>
      </tr>
      <tr>
        <th>Exhuma</th>
        <td>7.0</td>
        <td>66.9</td>
      </tr>
      <tr>
        <th>The First Omen</th>
        <td>6.7</td>
        <td>65.0</td>
      </tr>
      <tr>
        <th>Abigail</th>
        <td>6.7</td>
        <td>62.0</td>
      </tr>
      <tr>
        <th>Immaculate</th>
        <td>5.8</td>
        <td>57.0</td>
      </tr>
      <tr>
        <th>Sting</th>
        <td>5.7</td>
        <td>57.0</td>
      </tr>
      <tr>
        <th>The Strangers: Chapter 1</th>
        <td>5.0</td>
        <td>43.0</td>
      </tr>
      <tr>
        <th>Tarot</th>
        <td>4.9</td>
        <td>36.0</td>
      </tr>
    </tbody>
  </table>
</div>


```python
def genre_specific(df, genre):
    ...
```

```python
grader.check("q08_3")
```

## Finish Line 🏁

Congratulations! You're ready to submit Homework 2.

To submit your homework:

1. Select `Kernel -> Restart & Run All` to ensure that you have executed all cells, including the test cells.
2. Read through the notebook to make sure everything is fine and all tests passed.
3. Run the cell below to run all tests, and make sure that they all pass.
4. Download your notebook using `File -> Download as -> Notebook (.ipynb)`, then upload your notebook to Gradescope under "Homework 2".
5. Stick around while the Gradescope autograder grades your work. Make sure you see that all **public tests** have passed on Gradescope. **Remember that homeworks have hidden tests, which you will not see your scores on until a few days after the deadline!**
6. Check that you have a confirmation email from Gradescope and save it as proof of your submission.

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q01 (points: 4)
```
>>> callable(create_adjacency) and isinstance(create_adjacency(example_net), np.ndarray)
True
```
```
>>> example_net_output = create_adjacency(example_net)
>>> example_net_output.shape == (4, 4) and np.allclose(example_net_output, A)
True
```
```
>>> public_test_net = {1: [2, 3], 2: [1, 3], 3: [1, 2]}
>>> public_test_net_output = create_adjacency(public_test_net)
>>> public_test_net_output.shape == (3, 3) and np.allclose(public_test_net_output, np.array([[0, 0.5, 0.5], [0.5, 0, 0.5], [0.5, 0.5, 0]]))
True
```


Tests for q02_1 (points: 2)
```
>>> callable(compute_scores) and isinstance(compute_scores(A), np.ndarray)
True
```
```
>>> np.allclose(compute_scores(A), np.array([0.30769231, 0.38461538, 0.07692308, 0.23076923]))
True
```


Tests for q02_2 (points: 2)
```
>>> callable(pagerank) and isinstance(pagerank(A), list)
True
```
```
>>> np.allclose(pagerank(A), [2, 1, 4, 3])
True
```


Tests for q04 (points: 5)
```
>>> callable(prob_at_least_k_same_suit) and isinstance(prob_at_least_k_same_suit(4), float)
True
```
```
>>> call_one = prob_at_least_k_same_suit(4)
>>> call_two = prob_at_least_k_same_suit(4)
>>> call_one != call_two
True
```
failure_message: Make sure your function doesn't just return the same value every time!
```
>>> 0.022 <= prob_at_least_k_same_suit(4) <= 0.029
True
```
failure_message: Failing this but think you did your simulation correctly? Check your replacement strategy.


Tests for q05 (points: 5)
```
>>> callable(prob_all_one_jack) and isinstance(prob_all_one_jack(), float)
True
```
```
>>> call_one = prob_all_one_jack()
>>> call_two = prob_all_one_jack()
>>> call_one != call_two
True
```
failure_message: Make sure your function doesn't just return the same value every time!


Tests for q06_1 (points: 1)
```
>>> (isinstance(highest_rating_ever, float) or isinstance(highest_rating_ever, np.float64)) and highest_rating_ever in list(imdb['IMDb Rating'])
True
```


Tests for q06_2 (points: 2)
```
>>> isinstance(movie_with_highest_rating_ever, str) and movie_with_highest_rating_ever in list(imdb['Title'])
True
```


Tests for q06_3 (points: 1)
```
>>> (isinstance(num_comedy, int) or isinstance(num_comedy, np.int64)) and num_comedy < imdb.shape[0]
True
```


Tests for q06_4 (points: 2)
```
>>> isinstance(most_common_genre_2024, str) and most_common_genre_2024 in list(imdb['Genre'])
True
```


Tests for q06_5 (points: 3)
```
>>> (isinstance(prop_big_directors, float) or isinstance(prop_big_directors, np.float64)) and 0 < prop_big_directors < 1
True
```


Tests for q06_7 (points: 4)
```
>>> (isinstance(num_hours_family_friendly, float) or isinstance(num_hours_family_friendly, np.float64)) and num_hours_family_friendly < 48
True
```


Tests for q07_1 (points: 5)
```
>>> callable(extract_names) and isinstance(extract_names('Makoto ShinkaiClark Cheng'), list)
True
```
```
>>> extract_names('Makoto ShinkaiClark Cheng') == ['Makoto Shinkai', 'Clark Cheng']
True
```
```
>>> extract_names('Anya Taylor-JoyChris HemsworthTom Burke') == ['Anya Taylor-Joy', 'Chris Hemsworth', 'Tom Burke']
True
```


Tests for q07_2 (points: 2)
```
>>> callable(most_common) and isinstance(most_common(pd.Series([1, 2, 2])), list)
True
```
```
>>> most_common(pd.Series([1, 2, 2])) == [2]
True
```
```
>>> most_common(pd.Series([1, 2])) == [1, 2] or most_common(pd.Series([1, 2])) == [2, 1]
True
```


Tests for q07_6 (points: 2)
```
>>> (isinstance(average_harry_potter, float) or isinstance(average_harry_potter, np.float64)) and 3.9 < average_harry_potter < 9.2
True
```


Tests for q08_1 (points: 3)
```
>>> callable(longest_title) and isinstance(longest_title(imdb.head(5)), str)
True
```
```
>>> longest_title(imdb.head(5)) == '10 Things I Hate About You'
True
```
```
>>> longest_title(imdb.tail(5)) == "You Can't Run Forever"
True
```


Tests for q08_2 (points: 3)
```
>>> callable(metascore_to_rating_outlier) and isinstance(metascore_to_rating_outlier(imdb.head(5)), str)
True
```
```
>>> metascore_to_rating_outlier(imdb.head(5)) == 'A Quiet Place'
True
```
```
>>> metascore_to_rating_outlier(imdb.iloc[15:25]) == 'Barbie'
True
```


Tests for q08_3 (points: 3)
```
>>> callable(genre_specific) and isinstance(genre_specific(imdb, 'Horror'), pd.DataFrame)
True
```
```
>>> horror_output = genre_specific(imdb, 'Horror')
>>> horror_output.shape == (8, 2)
True
```
```
>>> horror_output = genre_specific(imdb, 'Horror')
>>> np.all(horror_output.index == ['Alien: Romulus', 'Exhuma', 'The First Omen', 'Abigail', 'Immaculate', 'Sting', 'The Strangers: Chapter 1', 'Tarot'])
True
```
